# Hello Python! 🐍

References:

[1] B Downey, A. (2012). *Think Python: How to Think Like a Computer Scientist-2e.*

[2] Gries, P., Campbell, J., & Montojo, J. (2017). *Practical programming: an introduction to computer science using Python 3.6.* Pragmatic Bookshelf.

[3] Matthes, E. (2023). *Python crash course: A hands-on, project-based introduction to programming.*References:

## 1 The Ceremonial `hello world!`

Let us know perform the ceremonial printing of `hello world!` using python!

### Approach 1

```
start
print the string "hello world!"
end
```

```python
print("hello world!")
```

Output:
```
hello world!
```

### Approach 2

```
start
set `message` as `"hello world!"`
print `message`
end
```

```python
message = "hello world!"
print(message)
```

Output:
```
hello world!
```

The difference is that in approach 2, we've set the *variable* `message` with the value `"hello world!"` then printed the `message` in our notebook.

<div class="alert alert-success">

**PEP 8: Assignment Statements**
    
Always surround these binary operators with a single space on either side: assignment (=), augmented assignment (+=, -= etc.), comparisons (==, <, >, !=, <>, <=, >=, in, not in, is, is not), Booleans (and, or, not).

Source: https://peps.python.org/pep-0008/#other-recommendations

</div>

When we are using variables in Python, there are some rules and guidelines.

1. Variable names can contain only letters, numbers, and underscores. They can start with a letter or an underscore, but not with a number. e.g., `message_1` is valid but `1_message` is not.
2. Spaces are not allowed in variable names, but underscores can be used to separate words in variable names. e.g.,` greeting_message` works, but `greeting message` will cause errors.
3. Avoid using Python keywords and function names as variable names; do not use words that Python has reserved for a particular purpose. e.g., `print`. See [Python Keywords](https://www.w3schools.com/python/python_ref_keywords.asp) for more examples.
4. Variable names should be short but descriptive.
5. Be careful when using lower case `l` and upper case `O`, because they can be confused with `1` and `0`.

<div class="alert alert-success">

**PEP 8: Function and Variable Names (Recommendation)**
    
Function names should be lowercase, with words separated by underscores as necessary to improve readability.

Variable names follow the same convention as function names.

mixedCase is allowed only in contexts where that’s already the prevailing style (e.g. threading.py), to retain backwards compatibility.

Source: https://peps.python.org/pep-0008/#function-and-variable-names

</div>

## 2 Strings

A *string* is a series of characters. We can define strings in python using single or double quotes.

### Defining strings

```python
"This is a string"
```

```python
'This is a string'
```

```python
'I told my friend, "My dream is to become a data science leader"'
```

```python
"The language 'Python' is named after Monty Python"
```

<div class="alert alert-success">

**PEP 8: String Quotes**
    
In Python, single-quoted strings and double-quoted strings are the same. This PEP does not make a recommendation for this. Pick a rule and stick to it. When a string contains single or double quote characters, however, use the other one to avoid backslashes in the string. It improves readability.

For triple-quoted strings, always use double quote characters to be consistent with the docstring convention in PEP 257.

Source: https://peps.python.org/pep-0008/#string-quotes

</div>

### String operations

Here are some other operations that you can do with strings:

```python
name = "Leo Lorenzo II"

print(name.title())
print(name.upper())
print(name.lower())
```

```python
help(str)
```

```python
"Leo" + " " + "Lorenzo"
```

```python
"1" + "6"
```

```python
'eo' in name
```

```python
'eo' * 10
```

```python
len('eo' * 10)
```

```python
min('Leo')
```

```python
max('data science')
```

```python
'data science'.index('e')
```

```python
'data science'.count('e')
```

```python
name
```

```python
name[0]
```

```python
name[-1]
```

```python
name[0:3]
```

```python
name[:3]
```

```python
name[0:-1]
```

```python
name[::-1]
```

```python
name.center(30)
```

```python
name.ljust(30)
```

```python
name.rjust(30)
```

### Using variables in strings

```python
first_name = "leo"
last_name = "lorenzo"
suffix = "ii"
full_name = f"{first_name} {last_name} {suffix}" # F-string

print(full_name)
```

```python
print("{} {} {}".format(first_name, last_name, suffix)) # format method
```

```python
print("%s %s %s"%(first_name, last_name, suffix)) # traditional
```

## 3 Numbers

Now, let's look at how numbers are defined in Python.

### Defining numbers

Numbers can be any of the following type:

1. `int` - Integers (a number that is not a fraction; a whole number)
2. `float` - Any number with a decimal point

```python
1
```

```python
1.
```

```python
type(1)
```

```python
type(1.)
```

When you divide any two numbers, the result is always a float

```python
1 / 1
```

But you can force the output to be of a specific type by using the `float` or `int` functions

```python
int(1 / 1)
```

To make numbers more readable, you can also use `_` to group digits.

```python
1000000000
```

```python
1_000_000_000
```

Or print them with a specific formatting

```python
large_number = 1_000_000_000
print(f"{large_number:,}")
```

### Number operations

Here are some arithmetic operators that we can use with numeric values

| Operator | Name |
| ------ | ----- |
| + | Addition |
| - | Subtraction |
| * | Multiplication |
| / | Division |
| % | Modulus |
| ** | Exponentiation |
| // | Floor division |

```python
9 - 3
```

```python
8 * 2.5
```

```python
9 / 2
```

```python
9 / -2
```

```python
9 // 2
```

```python
9 // -2
```

```python
9 % 2
```

```python
9.0 % 2
```

```python
9 % 2.0
```

```python
-9 % 2
```

```python
9 % 4
```

```python
-9 % 4
```

```python
6 - (9 % 6)
```

```python
-9 % 6
```

```python
4 - (9 % 4)
```

```python
9 / -2.0
```

```python
4 + 3*5
```

```python
(4 + 3)*5
```

But, wait, there's more!

| Operations | Name |
| ------ | ----- |
| abs(x) | Absolute value of x |
| int(x) | Converts x into an integer |
| float(x) | Converts x into floating point |
| complex(x, y) | Defines a complex number x + iy |
| c.conjugate() | Conjugates complex number c |
| divmod(x, y) | the pair (x // y, x % y) |
| pow(x, y) | x to the power of y |
| x ** y | x to the power of y |



```python
abs(-1)
```

```python
int(1.23)
```

```python
float('1.34')
```

```python
z = complex(1, 2)
z
```

```python
z.conjugate()
```

```python
divmod(11, 5)
```

```python
pow(2, 3)
```

```python
2**3
```

## 4 Hands-on Exercises

### Famous Quote

Find a famous quote from a famouse person you admire. Represent the famous person's name using a variable `famous_person`. Then compose your message and represent it with a new variable called `message`. Print your `message`.

The output should look like the following (including the quotation marks):

```
Albert Einstein once said, "A person who never made a mistake never tried anything new."
```

<!-- grade_id: cell-47268e81109fbb26, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

### Temperature converter

Given a temperature in variable `celsius_temp` in degrees celsius, convert this temperature value into the fahrenheit scale then represent it using a new variable called `fahrenheit_temp`. The conversion between degrees celsius and fahrenheit is shown by the equation below:


$$
F = \frac{9}{5} C + 32
$$

<!-- grade_id: cell-8fa7e97d208bdf8c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

### Measuring distances

Create a code that prints the Pythagorean (euclidean) `distance` between the two points. Represent the x and y coordinate of the first point as `x1` and `y1`, while set `x2` and `y2` the coordinate of the second point.

<!-- grade_id: cell-29aba1b592bc2566, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

### Time converter

Create a code that accepts a user input `time_secs` in seconds then displays this the value in minutes and seconds.

For example, if the elapsed time was `130` seconds, the output would be

```
2 minutes 10 seconds
```

<!-- grade_id: cell-490569c100e77829, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

```python
a = 4
```

```python
print('This is a line \n Another line')
print("other line")
```

Output:
```
This is a line 
 Another line
other line
```

```python

```

<!-- grade_id: cell-84f05c23c631eb26, points: 5, grade: true, solution: false -->
```python
assert True
```