# Loops and Conditionals

References:

[1] B Downey, A. (2012). *Think Python: How to Think Like a Computer Scientist-2e.*

[2] Gries, P., Campbell, J., & Montojo, J. (2017). *Practical programming: an introduction to computer science using Python 3.6.* Pragmatic Bookshelf.

[3] Matthes, E. (2023). *Python crash course: A hands-on, project-based introduction to programming.*References:

## 1 Conditionals

Programming often involves examining a set of conditions then deciding which actino to take based on those conditions. In the following section, we'll discuss the building blocks that we need to integrate conditional logic to our code.

### Conditional tests

First thing that we need are expressions that can be evaluated as `True` or `False`. In Python, these are expressed as *conditional tests*. Most of these tests compare the current value of a variable to a specific value of interest.

#### Checking for equality

For example, to check for equality between two values, we use the equality operation denoted by `==`. This operator returns `True` if the values on the left and right side of the operator match, and `False` if they don't match.

```python
1 == 2
```

Output:
```
False
```

```python
2 == 2
```

Output:
```
True
```

```python
car = 'toyota'
print(car == 'honda')
```

Output:
```
False
```

Note that when comparing string, testing for equality is case sensitive. This means that strings with different capitalization are not considered equal

```python
car = 'Toyota'
print(car == 'toyota')
```

Output:
```
False
```

What we can do to mediate this is to normalize the case of a variable's value just before doing the comparison.

```python
car.lower() == 'toyota'
```

Output:
```
True
```

#### Checking for Inequality

When we want to determine if two values are not equal, we can use the not equal operator. This is denoted by the `!=` operator.

```python
customer_order = 'ramen'
print(customer_order != 'pho')
```

Output:
```
True
```

#### Numerical comparisons

Testing numerical values is also straightforward, refer to the table below for the different comparison you can make and the corresponding operators.

| Operator | Description |
| -------- | ----------- |
| `==` | If the values of the operands are equal, then the condition becomes `True` | 
| `!=` | If the values of the operands are NOT equal, then this will return `True` |
| `>` | If the value of the left operand is greater than the value of right operand, this will return `True` |
| `<` | If the value of the left operand is less than the value of the right operand, this will return `False` |
| `>=` | If the value of the left operand is greater than or equal to the value of right operand, this will return `True` |
| `<=` | If the value of the left operand is less than or equal to the value of right operand, this will return `True` |

For example, you are trying to check the age of a particular person, you can do:

```python
age = 26
print("Age is equal to 18: {}".format(age == 18))
print("Age is NOT equal to 18: {}".format(age != 18))
print("Age is < 18: {}".format(age < 18))
print("Age is > 18: {}".format(age > 18))
print("Age is ≥ 18: {}".format(age >= 18))
print("Age is ≤ 18: {}".format(age <= 18))
```

Output:
```
Age is equal to 18: False
Age is NOT equal to 18: True
Age is < 18: False
Age is > 18: True
Age is ≥ 18: True
Age is ≤ 18: False
```

#### Checking multiple conditions

To check for multiple conditions, you can use the `and` and `or` operator. The `and` operator returns `True` if the operands are both `True`. Meanwhile, the `or` operator returns `True` if at least one of the operand is `True`.

```python
age = 26

print(age >= 18 and age < 30)
```

Output:
```
True
```

```python
age < 18 or age > 30
```

Output:
```
False
```

#### Boolean Expressions

Notice that the printed values for the above conditionals correspond to the value of either `True` or `False`. This kind of value is what we call a `Boolean value`.

```python
type(age < 18)
```

Output:
```
bool
```

Boolean values are often used to keep track of certain conditions such as whether a person is active or whether a person has certain permissions.

```python
is_active = True
can_view = True
can_upload = False
can_edit = False
```

#### Using the `not` keyword

The `not` keyword can also be used to negate any *boolean* value.

```python
not (age >= 18)
```

Output:
```
False
```

```python
age = 26

print(not(age > 30))
```

Output:
```
True
```

## 2 `if` Statements

Now we understand how to make conditional tests, we can start to integrate logic to our code.

### Simple `if` Statements

The simplest kind of `if` statement has one test and one action:

```python
if conditional_test:
    do something
```

```python
age = 17
if age >= 17:
    print("You are eligible to apply for a professional driver's license!")
```

Output:
```
You are eligible to apply for a professional driver's license!
```

### `if`-`else` statements

Oftentimes, you'll want to take one action when a conditional test passes and a different action in all other cases. We can use Python's `if`-`else` syntax to make this possible.

```python
if conditional_test:
    # Do something if `conditional_test` is True
else:
    # Do something if `conditional_test` is False
```

```python
age = 20
if age >= 17:
    print("You are eligible to apply for a professional driver's license!")
else:
    print("I'm sorry, you cannot apply for a professional driver's license.")
```

Output:
```
You are eligible to apply for a professional driver's license!
```

### `if`-`elif`-`else` chain

In other times, you'll need to test more than two possible situations. To evaluate this, you can use the Python's `if-elif-else` syntax. Python will execute only one block in the `if-elif-else` chain. It runs each conditional test in order until one passes. When a test passes, the code following that test is executed and Python skips the rest of the tests.

A real world example of this is typhoon classification.

```python
wind_speed = 120 # wind speed in kilometers per hour

if wind_speed < 62:
    classification = "TROPICAL DEPRESSION (TD)"
elif wind_speed < 88:
    classification = "TROPICAL STORM (TS)"
elif wind_speed < 117:
    classification = "SEVERE TROPICAL STORM (STS)"
elif wind_speed < 184:
    classification = "TYPHOON (TY)"
else:
    classification = "SUPER TYPHOON (STY)"

print(f"The classification of the typhoon is {classification}")
```

Output:
```
The classification of the typhoon is TYPHOON (TY)
```

Note that Python does not require an `else` block at the end of the `if-elif` chain. Sometimes an else block can be useful, sometimes it is clearer to use an additional `elif` statement that catches the specific condition of interest.

## 3 Loops

Up until know, if we want to execute a code two hundred times, you would need to run that specific code block/cell two hundred times. For this section, we will introduce another fundamental kind of control flow: repetition. We'll see how we can write an instruction once and use loops to repeat that code the desired number of times.

There are two ways to use loops: the first one is using an iterable which allows us to process items in a collection of data; the second one is to use conditionals as trigger when to stop executing the code in the loop. We'll consider both of these in the next subsections.

### Looping through an entire collection

We'll often would want to run through different elements in an iterable, performing the same task in each item. As an example, consider a task wherein we are to check whether a particular letter in a string is in uppercase. We can use Python's `for` loop syntax to process each character, without having to write one statement per element.

```python
message = 'I am a Data Science Leader'
for ch in message:
    if ch.isupper():
        print(ch)
```

Output:
```
I
D
S
L
```

The general form of a `for` loop is as follows:

```python
for <<loop_variable>> in <<list>>:
    <<block>>
```

The way a `for` loop is executed is as follows:
1. The loop variable is assigned the first item in the list, and the loop block is executed.
2. The loop variable is then assigned to the second item in the list and the loop body is executed again
3. This is repeated until the loop variable is assigned the last item in the list and the loop body is executed one last time.

Always remember to indent the line after the `for` statement in a loop. If you forget, Python will remind you by throwing off and error.

#### Using the `range()` function

Python has a built-in function `range()` that can be useful to generate a series of numbers.

```python
help(range)
```

Output:
```
Help on class range in module builtins:

class range(object)
 |  range(stop) -> range object
 |  range(start, stop[, step]) -> range object
 |  
 |  Return an object that produces a sequence of integers from start (inclusive)
 |  to stop (exclusive) by step.  range(i, j) produces i, i+1, i+2, ..., j-1.
 |  start defaults to 0, and stop is omitted!  range(4) produces 0, 1, 2, 3.
 |  These are exactly the valid indices for a list of 4 elements.
 |  When step is given, it specifies the increment (or decrement).
 |  
 |  Methods defined here:
 |  
 |  __bool__(self, /)
 |      self != 0
 |  
 |  __contains__(self, key, /)
 |      Return key in self.
 |  
 |  __eq__(self, value, /)
 |      Return self==value.
 |  
 |  __ge__(self, value, /)
 |      Return self>=value.
 |  
 |  __getattribute__(self, name, /)
 |      Return getattr(self, name).
 |  
 |  __getitem__(self, key, /)
 |      Return self[key].
 |  
 |  __gt__(self, value, /)
 |      Return self>value.
 |  
 |  __hash__(self, /)
 |
```

You can use the `range()` in two ways:

First is by using an integer input

```python
for value in range(5):
    print(value)
```

Output:
```
0
1
2
3
4
```

The second way is to specify the starting value, ending value, and optionally, the increment or decrement value.

```python
for value in range(1, 5):
    print(value)
```

Output:
```
1
2
3
4
```

### Looping using conditionals

`for` loops are useful if you know how many iterations of the loop you need. In some situations, you might not know in advance how many loop iterations you have to execute. In these cases, we can use a `while` loop. The general form of a `while` loop is as follows:

```python
while <<conditional_expression>>:
    <<block>>
```

Here's an example:

```python
rabbits = 3
while rabbits > 0:
    print(f"Number of rabbits: {rabbits}")
    rabbits -= 1
```

Output:
```
Number of rabbits: 3
Number of rabbits: 2
Number of rabbits: 1
```

Another more useful example is a simple exponential growth model. Consider the following bacteria population model:

$$
P(t + 1) = P(t) + rP(t)
$$

where $P(t)$ is the population size at time $t$ in minutes and $r$ is the growth rate. Let's create a code that determines how long it takes for a bacteria to double their numbers.

```python
time = 0
initial_population = 1_000 # Initial population
growth_rate = 0.21 # 21% growth per minute
population = initial_population

# We terminate the loop after twice the initial population
while population < 2*initial_population:
    population += growth_rate*population
    print(f'Current population: {round(population)}')
    time += 1

print(f"\nIt took {time} minutes for the bacteria to double.")
print(f"The final population was {round(population)} bacteria.")
```

Output:
```
Current population: 1210
Current population: 1464
Current population: 1772
Current population: 2144

It took 4 minutes for the bacteria to double.
The final population was 2144 bacteria.
```

### Controlling Loops Using `break` and `continue`

As a rule, `for` and `while` loops execute all the statements in their body on each iteration. However, sometimes it is useful to break that rule. Python provides two ways of controlling the iteratin of a loop: `break` - which terminates the execution of the loop immediately, and `continue`, which skips ahead to the next iteration.

For example, let's look at a simple worker that accomplishes `1` unit of work every hour. Let's see what happens to our loops when we use some `break` and `continue` statements.

```python
# Case when we don't use `continue` or `break`
work_done = 0 # Initialize the work done of a worker

for time in range(24): # Time in hours
    work_done += 1

print(f"Total work done: {work_done} units")
```

Output:
```
Total work done: 24 units
```

```python
# Case when we use break
work_done = 0

for time in range(24):
    if time == 8:
        break
    else:
        work_done += 1

print(f"Total work done: {work_done} units")
```

Output:
```
Total work done: 8 units
```

```python
# Case when we use continue
work_done = 0

for time in range(24):
    if time == 8:
        continue
    else:
        work_done += 1

print(f"Total work done: {work_done} units")
```

Output:
```
Total work done: 23 units
```

```python
work_done = 0

for i, time in enumerate(range(24)):
    if time == 8:
        continue
    else:
        work_done += 1
    print(f"Time: {i+1}; Work done: {work_done}")
```

Output:
```
Time: 1; Work done: 1
Time: 2; Work done: 2
Time: 3; Work done: 3
Time: 4; Work done: 4
Time: 5; Work done: 5
Time: 6; Work done: 6
Time: 7; Work done: 7
Time: 8; Work done: 8
Time: 10; Work done: 9
Time: 11; Work done: 10
Time: 12; Work done: 11
Time: 13; Work done: 12
Time: 14; Work done: 13
Time: 15; Work done: 14
Time: 16; Work done: 15
Time: 17; Work done: 16
Time: 18; Work done: 17
Time: 19; Work done: 18
Time: 20; Work done: 19
Time: 21; Work done: 20
Time: 22; Work done: 21
Time: 23; Work done: 22
Time: 24; Work done: 23
```

## 4 Hands-on Exercises

### Stages of Life

Create a code that determines a person's stage of life given their `age`. The code will output a value according to the following conditions:
- If a person's `age` is less than 2 years old, that person is a `baby`.
- If a person's `age` is at least 2 years old but less than 4, that person is a `toddler`
- If a person's `age` is at least 4 years old but less than 13, that person is a `kid`.
- If a person's `age` is at least 13 years old but less than 20, that person is a `teenager`.
- If a person's `age` is at least 20 years old but less than 65, that person is an `adult`.
- If a person's `age` is 65 years or older, that person is an `elder`.

<!-- grade_id: cell-05624bd67a1bfbe3, points: 0, grade: false, solution: true -->
```python
age = 28
### BEGIN SOLUTION

### END SOLUTION
```

Output:
```
adult
```

### Generating Triangles

Using `for` loops, print a right triangle `T` on the screen where the triangle is one character wide at its narrow point and seven characters wide at its widest point.

```
T
TT
TTT
TTTT
TTTTT
TTTTTT
TTTTTTT
```

<!-- grade_id: cell-0ca2d7bedc42e90e, points: 0, grade: false, solution: true -->
```python
testStr = "ae"
testStr2 = "io"

print(testStr + testStr2)
print(testStr*3)
```

Output:
```
TypeError                                 Traceback (most recent call last)
Cell In[3], line 4
      1 testStr = "ae"
      2 testStr2 = "io"
----> 4 print(testStr - testStr2)
      5 print(testStr*3)

TypeError: unsupported operand type(s) for -: 'str' and 'str'
```

Now try to print the triangle described in the previous exercise with its hypothenuse on the left side.

```
      T
     TT
    TTT
   TTTT
  TTTTT
 TTTTTT
TTTTTTT
```

<!-- grade_id: cell-c6ecc968c2621c87, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
for i in range(7):
    
### END SOLUTION
```

Output:
```
T
     TT
    TTT
   TTTT
  TTTTT
 TTTTTT
TTTTTTT
```

### Fizz Buzz Fuzz

Create a code that checks the value of an integer `n`, then prints a string output according to the following conditions:

1. Print `'Fizz'` if `n` is divisble by 3.
2. Print `'Buzz'` if `n` is divisible by 5.
3. Print `'Fuzz'` if `n` is divisble by 15.
4. Print the value `n` as a string if none of the above conditions are true.

<!-- grade_id: cell-7e72281ca29f45ec, points: 0, grade: false, solution: true -->
```python
n = 25
### BEGIN SOLUTION

### END SOLUTION
```

Output:
```
Buzz
```

### Them rats!

Variables `rat_1_weight` and `rat_2_weight` contain the weights of two rats at the beginning of the experiment. Variables `rat_1_rate` and `rat_2_rate` are the rate that the rats' weights are expected to increase each week. (for example, 5 percent per week for rat 1, and 2 percent per week for rat 2)

```python
rat_1_weight = 200
rat_2_weight = 200
rat_1_rate = 0.05
rat_2_rate = 0.02
```

1. Using a `while` loop, calculate how many `weeks` it would take for the weight of the first rat to become 25 percent heavier than it was originally.

<!-- grade_id: cell-74fc5b8b9ea5df29, points: 0, grade: false, solution: true -->
```python
weeks = 0
### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-22a8245129ed829e, points: 1, grade: true, solution: false -->
```python
assert isinstance(weeks, int)
### BEGIN HIDDEN TESTS

### END HIDDEN TESTS
```

2. Using a `while` loop, calculate how many weeks it would take for rat 1 to be 10 percent heavier than rat 2.

<!-- grade_id: cell-fe574e6932fe3f02, points: 0, grade: false, solution: true -->
```python
weeks = 0
### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-12cdd8273c58be62, points: 1, grade: true, solution: false -->
```python
assert isinstance(weeks, int)
### BEGIN HIDDEN TESTS

### END HIDDEN TESTS
```