# Storing and Processing Collections of Data

References:

[1] B Downey, A. (2012). *Think Python: How to Think Like a Computer Scientist-2e.*

[2] Gries, P., Campbell, J., & Montojo, J. (2017). *Practical programming: an introduction to computer science using Python 3.6.* Pragmatic Bookshelf.

[3] Matthes, E. (2023). *Python crash course: A hands-on, project-based introduction to programming.*

Up until this point, we have seen several data types in python such as strings, functions, and other numerical data types. In this notebook, we introduce lists and dictionaries. Which allow you to store sets of information in one place.

Both of these types are very powerful features readily accessible to new programmers, and they help tie together many important concepts in programming.

## 1 Lists

A *list* is a collection of items in a particular order. We can put anything inside a list and the items in the list don't have to be related in any particular way.

In Python, to define a list we use the square brackets `[]` and individual elements in the list are separated by commas. Since lists usually contain more than one element, it's a good idea to make the name of the list plural, such as `letters`, `digits`, or `names`.

**Example**: A list that contains a few kinds of bicycles

```python
bicycles = ['trek', 'cannondale', 'redline', 'specialized']
print(bicycles)
```

Output:
```
['trek', 'cannondale', 'redline', 'specialized']
```

### Accessing elements in a list

Lists are ordered collections, thus, we can access any element by telling the Python the position, or *index*, of the item desired.

*Note: always remember that Python considers the first item in a list to be at position 0 not position 1*

For example, to pull out the first bicycle in our list `bicycles` we do:

```python
print(bicycles[0])
```

Output:
```
trek
```

Python also has a special syntax for accessing the last element in the list. We can use the index `-1` to get the last item in the list.

The index `-2` returns the second item from the end of the list, the index `-3` returns the third, and so forth.

```python
print(bicycles[-1])
```

Output:
```
specialized
```

You can also access a subset of a list by *slicing* it. To make a slice, specify the first and last elements you want to work with. Python stops one item before the second index you specify.

```python
print(bicycles)
```

Output:
```
['trek', 'cannondale', 'redline', 'specialized']
```

```python
print(bicycles[0:2])
```

Output:
```
['trek', 'cannondale']
```

### Changing, Adding, and Removing Elements

To modify an item in the list, we can access first the element of a list then perform an assignment operation to the new value you want the item to have.

```python
motorcycles = ['honda', 'yamaha', 'suzuki']
print(motorcycles)
```

Output:
```
['honda', 'yamaha', 'suzuki']
```

```python
motorcycles[0] = 'ducati'
print(motorcycles)
```

Output:
```
['ducati', 'yamaha', 'suzuki']
```

Aside from modifying the elements, we can also *append* a new item to the list.

```python
motorcycles.append('honda')
print(motorcycles)
```

Output:
```
['ducati', 'yamaha', 'suzuki', 'honda']
```

Another way to do this is to extend the list using another list

```python
motorcycles.extend(['kawasaki', 'triumph'])
print(motorcycles)
```

Output:
```
['ducati', 'yamaha', 'suzuki', 'honda', 'kawasaki', 'triumph']
```

You can also insert a specific element at a specific location on a list using the `insert` method. Here, we specify the element to be inserted and the specific index in which we wish to place it.

```python
motorcycles.insert(3, 'bmw')
print(motorcycles)
```

Output:
```
['ducati', 'yamaha', 'suzuki', 'bmw', 'honda', 'kawasaki', 'triumph']
```

To remove an element, we can use the `del` statement to delete a specific element in the list.

```python
del motorcycles[1]
print(motorcycles)
```

Output:
```
['ducati', 'suzuki', 'bmw', 'honda', 'kawasaki', 'triumph']
```

Another way to do this is the `pop` method. The difference is that it returns the value of the removed item in addition to removing it.

```python
print(motorcycles)
```

Output:
```
['ducati', 'suzuki', 'bmw', 'honda', 'kawasaki', 'triumph']
```

```python
popped_motorcycle = motorcycles.pop(-1)
print(motorcycles)
print(popped_motorcycle)
```

Output:
```
['ducati', 'suzuki', 'bmw', 'honda', 'kawasaki']
triumph
```

### List Operations

Aside from changing or modifying the elements of a single list, we can also also perform operations such as `+` and `*` on several lists.

```python
motorcycles + bicycles
```

Output:
```
['ducati',
 'suzuki',
 'bmw',
 'honda',
 'kawasaki',
 'trek',
 'cannondale',
 'redline',
 'specialized']
```

```python
motorcycles * 3
```

Output:
```
['ducati',
 'suzuki',
 'bmw',
 'honda',
 'kawasaki',
 'ducati',
 'suzuki',
 'bmw',
 'honda',
 'kawasaki',
 'ducati',
 'suzuki',
 'bmw',
 'honda',
 'kawasaki']
```

To get the number of elements in a list we can also use the built-in function `len`.

```python
print(len(motorcycles))
```

Output:
```
5
```

### Organizing a List

If we want to permanently sort a list, we use its `sort` method.

```python
print(motorcycles)
```

Output:
```
['ducati', 'suzuki', 'bmw', 'honda', 'kawasaki']
```

```python
motorcycles.sort()
print(motorcycles)
```

Output:
```
['bmw', 'ducati', 'honda', 'kawasaki', 'suzuki']
```

By specifying the `reverse` parameter as `True`, we can sort the items in descending order. 

```python
motorcycles.sort(reverse=True)
print(motorcycles)
```

Output:
```
['suzuki', 'kawasaki', 'honda', 'ducati', 'bmw']
```

You can also sort a list temporarily using the `sorted` function.

```python
motorcycles = ['ducati', 'suzuki', 'bmw', 'honda', 'kawasaki']

print(f"Original order: {motorcycles}")
print(f"Sorted order: {sorted(motorcycles)}")
print(f"After sorted operation: {motorcycles}")
```

Output:
```
Original order: ['ducati', 'suzuki', 'bmw', 'honda', 'kawasaki']
Sorted order: ['bmw', 'ducati', 'honda', 'kawasaki', 'suzuki']
After sorted operation: ['ducati', 'suzuki', 'bmw', 'honda', 'kawasaki']
```

### Processing Items of a List

There are various reasons as to why we want to store a set of numbers. For example, keeping track of the population of an endangered animal is an important aspect of wildlife preservation. Storing data in a list will also help in succeeding data analysis, such as the computation of simple statistics such as mean, median, or mode.

There are a few functions that will be helpful when working with lists of numbers.

```python
digits = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
print(min(digits))
print(max(digits))
print(sum(digits))
```

Output:
```
1
10
55
```

You can also iterate through each elements of a list using a `for` loop to process each element. For example, if we want to get the square of a series of numbers we can:

```python
squares = []
for value in digits:
    square = value ** 2
    squares.append(square)

print(squares)
```

Output:
```
[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
```

You'll notice that you'll be doing a lot of these kind of operation - transforming the items on a list then storing them on a new dictionary. For those cases, you can simplify your code greatly by using *list comprehensions*. Verify that the code above is equivalent to the one below.

```python
squares = [value ** 2 for value in digits]

print(squares)
```

Output:
```
[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
```

Sometimes you might also want to add a conditional before executing the transformation. In cases like this, you add a conditional after the required `for` loop statement.

```python
even_squares = []
for value in digits:
    if not value % 2:
        even_squares.append(value ** 2)

print(even_squares)
```

Output:
```
[4, 16, 36, 64, 100]
```

Verify that the above code is equivalent to:

```python
even_squares = [value ** 2 for value in digits if not value % 2]

print(even_squares)
```

Output:
```
[4, 16, 36, 64, 100]
```

### Lists and Strings

We can also convert strings to list and vice-versa.

```python
a_string = 'a string is a sequence of characters'
list(a_string)
```

Output:
```
['a',
 ' ',
 's',
 't',
 'r',
 'i',
 'n',
 'g',
 ' ',
 'i',
 's',
 ' ',
 'a',
 ' ',
 's',
 'e',
 'q',
 'u',
 'e',
 'n',
 'c',
 'e',
 ' ',
 'o',
 'f',
 ' ',
 'c',
 'h',
 'a',
 'r',
 'a',
 'c',
 't',
 'e',
 'r',
 's']
```

If we want more controllability on how we create the list from the string, we can use the `split` method of a string which also requires us to specify a delimiter.

```python
a_list = a_string.split(' ')
a_list
```

Output:
```
['a', 'string', 'is', 'a', 'sequence', 'of', 'characters']
```

You can reverse this process by using the `join` method. Here, we specify the delimiter as a string then place in the argument the list that we wish to join.

```python
print('\n'.join(a_list))
```

Output:
```
a
string
is
a
sequence
of
characters
```

## 2 Tuples

Lists are ideal for storing a colection of items that can change throughout the life of a program. However, sometimes you'll want to create a list of items that cannot change. Python refers to values that cannot change as *immutable*, and an immutable list is a tuple.

### Defining a tuple

```python
coordinate = (200, 50)
print(coordinate[0])
print(coordinate[1])
```

Output:
```
200
50
```

Tuples are immutable!

```python
# coordinate[0] = 100
```

### Tuples as `return` values

Because of this property, tuples are a natural data type that Python uses when functions return multiple values

```python
def min_max(sequence):
    return min(sequence), max(sequence)
```

```python
min_max(digits)
```

Output:
```
(1, 10)
```

```python
type(min_max(digits))
```

Output:
```
tuple
```

### Variable-length Argument Tuples

Tuples are also used as the container when specifying arbitrary length positional arguments

```python
def print_all(*pargs):
    """Print an arbitrary number of positional arguments as a tuple"""
    print(pargs)
```

```python
print_all(1, 2.0, '3')
```

Output:
```
(1, 2.0, '3')
```

```python
def print_all(*pargs):
    """Print an arbitrary number of positional arguments in a line"""
    print(*pargs)
```

```python
print_all(1, 2.0, '3')
```

Output:
```
1 2.0 3
```

```python
print(1, 2.0, '3')
```

Output:
```
1 2.0 3
```

### Lists and Tuples

If we are given two or more sequences, we can use the `zip` function to create an iterable that interleaves the values of the two sequence then returns the resulting sequence as an iterator.

```python
names = ['Leo', 'Pat', 'Gil']
ages = [32, 27, 30]

for name, age in zip(names, ages):
    print(f'{name}: {age}')
```

Output:
```
Leo: 32
Pat: 27
Gil: 30
```

Notice that if we inspect the `zip` object, it doesn't print its contents.

```python
zipped_sequence = zip(names, ages)
zipped_sequence
```

Output:
```
<zip at 0x10474a6c0>
```

This is because the `zip` object is a *lazy* iterators. *lazy* objects only evaluate their values until the value is needed. Once completely iterated, lazy objects also get exhausted - meaning they don't return any value.

```python
list(zipped_sequence)
```

Output:
```
[('Leo', 32), ('Pat', 27), ('Gil', 30)]
```

```python
list(zipped_sequence)
```

Output:
```
[]
```

To unzip a zipped sequence, we can scatter or unpack the zipped sequence on a zip function call.

```python
zipped_sequence = list(zip(names, ages))
zipped_sequence
```

Output:
```
[('Leo', 32), ('Pat', 27), ('Gil', 30)]
```

```python
print(tuple(zip(*zipped_sequence)))
```

Output:
```
(('Leo', 'Pat', 'Gil'), (32, 27, 30))
```

In this case, we can unpack the sequence then save it in a set of variables.

```python
same_names, same_ages = zip(*zipped_sequence)
print(same_names)
print(same_ages)
```

Output:
```
('Leo', 'Pat', 'Gil')
(32, 27, 30)
```

## 3 Dictionaries

Unlike list which stores a collection of data which are not necessarily related with each other, *dictionaries* enables us to store a collection of data with some related information. For example, we can create a dictionary representing information about a pokemon then store as much information as we want about that pokemon such as their name, type, height, weight, abilities, and base experience.

To define a dictionary, we use curly brackets `{}` then use `:` to denote a correspondence between information. Elements in the dictionary are separate by `,` just like lists.

```python
my_pokemon = {
    'Name': 'Evee', 'Type': 'Normal', 'Height': 0.30, 'Weight': 6.5,
    'Abilities': ['Run Away', 'Adaptability'], 'Base Exp': 65
}
```

<div class="alert alert-success">

**PEP 8: Lining up spaces in between braces/brackets/parenthesis**
    
The closing brace/bracket/parenthesis on multiline constructs may either line up under the first non-whitespace character of the last line of list or or it may be lined up under the first character of the line that starts the multiline construct.

Source: https://peps.python.org/pep-0008/#indentation

</div>

A dictionary is a collection of *key-value* pairs. Each *key* is connected to a value, and you can use a key to access the value associated with that key. The *value* of a particular *key* can be of any type, however, *keys* should be types that are *immutable*. Meaning, their values cannot be change by assignment.

### Accessing values in a dictionary

To get the value associated with a key, we give the name of the dictionary then place the key inside the square brackes after it.

```python
my_pokemon['Name']
```

Output:
```
'Evee'
```

We can also use the `get()` method to access values. This is much safer than just indexing a particular key since you can specify a default output if the key does not exist in the dictionary.

```python
my_pokemon.get('Name')
```

Output:
```
'Evee'
```

```python
# my_pokemon['Favorite Food']
```

```python
my_pokemon.get('Favorite Food', 'Pizza')
```

Output:
```
'Pizza'
```

The `get` method is essentially equivalent to the following function:

```python
def get_method(a_dict, key, default_value):
    if key in a_dict:
        return a_dict[key]
    else:
        return default_value
```

```python
get_method(my_pokemon, 'Favorite Food', 'Pizza')
```

Output:
```
'Pizza'
```

We can also obtain the same effect using the `setdefault()` method. However, when using `setdefault()`, we also insert the key with value of default if the key is not in the dictionary.

```python
my_pokemon
```

Output:
```
{'Name': 'Evee',
 'Type': 'Normal',
 'Height': 0.3,
 'Weight': 6.5,
 'Abilities': ['Run Away', 'Adaptability'],
 'Base Exp': 65}
```

```python
my_pokemon.setdefault('Favorite Food', 'Pizza')
```

Output:
```
'Pizza'
```

```python
my_pokemon
```

Output:
```
{'Name': 'Evee',
 'Type': 'Normal',
 'Height': 0.3,
 'Weight': 6.5,
 'Abilities': ['Run Away', 'Adaptability'],
 'Base Exp': 65,
 'Favorite Food': 'Pizza'}
```

### Changing, Adding, and Removing Elements

To change the value of a specific key pair, we can perform an assignment operation.

```python
my_pokemon['Name'] = 'Jolteon'
my_pokemon['Type'] = 'Electric'
print(my_pokemon)
```

Output:
```
{'Name': 'Jolteon', 'Type': 'Electric', 'Height': 0.3, 'Weight': 6.5, 'Abilities': ['Run Away', 'Adaptability'], 'Base Exp': 65, 'Favorite Food': 'Pizza'}
```

To add an element to our dictionary, we can perform an assignment to a key with no corresponding value.

```python
my_pokemon['Species'] = 'Lightning Pokemon'
print(my_pokemon)
```

Output:
```
{'Name': 'Jolteon', 'Type': 'Electric', 'Height': 0.3, 'Weight': 6.5, 'Abilities': ['Run Away', 'Adaptability'], 'Base Exp': 65, 'Favorite Food': 'Pizza', 'Species': 'Lightning Pokemon'}
```

To remove a piece of information, we use the `del` operator.

```python
del my_pokemon['Favorite Food']
print(my_pokemon)
```

Output:
```
{'Name': 'Jolteon', 'Type': 'Electric', 'Height': 0.3, 'Weight': 6.5, 'Abilities': ['Run Away', 'Adaptability'], 'Base Exp': 65, 'Species': 'Lightning Pokemon'}
```

We can also add contents from another dictionary using the `update` method

```python
my_pokemon.update({'EV yield': '2 Speed', 'Base Friendship': 50,
                   'Growth Rate': 'Medium Fast', 'Height': 0.8,
                   'Weight': 24.5})
```

```python
my_pokemon
```

Output:
```
{'Name': 'Jolteon',
 'Type': 'Electric',
 'Height': 0.8,
 'Weight': 24.5,
 'Abilities': ['Run Away', 'Adaptability'],
 'Base Exp': 65,
 'Species': 'Lightning Pokemon',
 'EV yield': '2 Speed',
 'Base Friendship': 50,
 'Growth Rate': 'Medium Fast'}
```

<div class="alert alert-success">

**PEP 8: Continuation Lines**
    
Continuation lines should align wrapped elements either vertically using Python’s implicit line joining inside parentheses, brackets and braces, or using a hanging indent.

Source: https://peps.python.org/pep-0008/#indentation

</div>

### Processing Items in a Dictionary

We can also process items in a dictionary using a `for` loop in conjunction with the `dict`'s `item()` method.

```python
for key, value in my_pokemon.items():
    print(f"{key}: {value}")
```

Output:
```
Name: Jolteon
Type: Electric
Height: 0.8
Weight: 24.5
Abilities: ['Run Away', 'Adaptability']
Base Exp: 65
Species: Lightning Pokemon
EV yield: 2 Speed
Base Friendship: 50
Growth Rate: Medium Fast
```

### Variable-length Keyword Arguments

When we place a `**` before a parameter name, it allows the function to **gather** an arbitrary number of keyword arguments into a corresponding dictionary.

```python
def print_pokemon(**kwargs):
    for key, value in kwargs.items():
        print(f"{key}: {value}")
```

```python
print_pokemon(Name='Evee', Height=0.30, Weight=6.5)
```

Output:
```
Name: Evee
Height: 0.3
Weight: 6.5
```

If we place a `**` before a dictionary during a function call, it **unpacks** or **scatters** the dictionary into the function's keyword parameters.

```python
print_pokemon(**my_pokemon)
```

Output:
```
Name: Jolteon
Type: Electric
Height: 0.8
Weight: 24.5
Abilities: ['Run Away', 'Adaptability']
Base Exp: 65
Species: Lightning Pokemon
EV yield: 2 Speed
Base Friendship: 50
Growth Rate: Medium Fast
```

## 4 Sets

For cases where we want to operate on collections with unique values, we can define a **set** object which is an unordered collection of unique items. Sets behave like a collection of dictionary keys with no values. As such, we can only store immutable objects inside the sets.

For example, here we define a set containing the vowel letters:

```python
vowels = {'a', 'e', 'i', 'o', 'u'}
vowels
```

Output:
```
{'a', 'e', 'i', 'o', 'u'}
```

Notice that if we include repeated items inside a set definition, duplicates are ignored:

```python
vowels = {'a', 'o', 'o', 'i', 'e', 'u', 'a'}
vowels
```

Output:
```
{'a', 'e', 'i', 'o', 'u'}
```

### Set operations

In mathematics, set operations include union, intersection, addition, and difference. In Python, we can implement this using methods or through the use of operators:

```python
vowels.add('y')
vowels
```

Output:
```
{'a', 'e', 'i', 'o', 'u', 'y'}
```

```python
ten = set(range(10))
lows = set(range(5))
evens = set([i for i in range(10) if not i % 2])
print(ten)
print(lows)
print(evens)
```

Output:
```
{0, 1, 2, 3, 4, 5, 6, 7, 8, 9}
{0, 1, 2, 3, 4}
{0, 2, 4, 6, 8}
```

```python
lows.difference(evens)
```

Output:
```
{1, 3}
```

```python
lows.intersection(evens)
```

Output:
```
{0, 2, 4}
```

```python
lows & evens
```

Output:
```
{0, 2, 4}
```

```python
lows.union(evens)
```

Output:
```
{0, 1, 2, 3, 4, 6, 8}
```

```python
lows | evens
```

Output:
```
{0, 1, 2, 3, 4, 6, 8}
```

```python
lows.issubset(ten)
```

Output:
```
True
```

```python
lows.issuperset(evens)
```

Output:
```
False
```

```python
lows.remove(0)
lows
```

Output:
```
{1, 2, 3, 4}
```

```python
lows.symmetric_difference(evens)
```

Output:
```
{0, 1, 3, 6, 8}
```

```python
lows.union(evens).difference(lows.intersection(evens))
```

Output:
```
{0, 1, 3, 6, 8}
```

```python
(lows | evens) - (lows & evens)
```

Output:
```
{0, 1, 3, 6, 8}
```

## 5 Other Collections

From the four container data types that we've covered, we can create representations to about any type of form of collections of data. However, through the use of specialized container types, we can further optimize our code. Here are few useful container types that will come in handy for specific tasks

### [collections](https://docs.python.org/3/library/collections.html) module

```python
from collections import Counter
```

```python
name = "Leodegario Lorenzo II"
char_count = Counter(name.lower())

print(char_count)
```

Output:
```
Counter({'o': 4, 'e': 3, 'i': 3, 'l': 2, 'r': 2, ' ': 2, 'd': 1, 'g': 1, 'a': 1, 'n': 1, 'z': 1})
```

```python
char_count.most_common(5)
```

Output:
```
[('o', 4), ('e', 3), ('i', 3), ('l', 2), ('r', 2)]
```

```python
from collections import defaultdict
```

```python
d_dict = defaultdict(lambda : None)
d_dict['name'] = 'Leo'
```

```python
d_dict
```

Output:
```
defaultdict(<function __main__.<lambda>()>, {'name': 'Leo'})
```

```python
d_dict['favorite_food']
```

### [itertools](https://docs.python.org/3/library/itertools.html) module

```python
import itertools
```

```python
itertools.chain('ABC', 'DEF')
```

Output:
```
<itertools.chain at 0x10476cf10>
```

```python
a_sequence = 'ABC'
b_sequence = 'DEF'
list(itertools.chain(a_sequence, b_sequence))
```

Output:
```
['A', 'B', 'C', 'D', 'E', 'F']
```

```python
list(itertools.product(a_sequence, b_sequence))
```

Output:
```
[('A', 'D'),
 ('A', 'E'),
 ('A', 'F'),
 ('B', 'D'),
 ('B', 'E'),
 ('B', 'F'),
 ('C', 'D'),
 ('C', 'E'),
 ('C', 'F')]
```

```python
list(itertools.permutations(a_sequence, 2))
```

Output:
```
[('A', 'B'), ('A', 'C'), ('B', 'A'), ('B', 'C'), ('C', 'A'), ('C', 'B')]
```

```python
list(itertools.combinations(a_sequence, 2))
```

Output:
```
[('A', 'B'), ('A', 'C'), ('B', 'C')]
```

## 6 Hands-on Exercises

### Draw a card

Create a function `draw_card()` which simulates you to randomly drawing a card in a standard 52-deck. The function should not accept any parameters but should return a string representing the card that was drawn, including its suit and rank.

Example output:

```
>>> draw_card()
'Jack of Spades'
```

*Hint: Use the `choice` method of the `random` library*

<!-- grade_id: cell-10fd12b092888dbc, points: 0, grade: false, solution: true -->
```python
import random

def draw_card():
    ### YOUR CODE HERE
    raise NotImplementedError
```

### Anagrams

Two words are anagrams if you can rearrange the letters from one to spell the other. Write a function named `is_anagram` that takes in two strings and returns `True` if they are anagrams, `False` otherwise

<!-- grade_id: cell-64d5996d7b57f2c8, points: 0, grade: false, solution: true -->
```python
def is_anagram(word1, word2):
    ### YOUR CODE HERE
    raise NotImplementedError
```

### Roman Numerals

Create a function `roman_to_integer()` that takes in a roman numeral string `roman_numeral` then outputs the corresponding integer value of the `roman_numeral`

<!-- grade_id: cell-04346a8c16bd772a, points: 0, grade: false, solution: true -->
```python
def roman_to_integer(roman_numeral):
    ### YOUR CODE HERE
    raise NotImplementedError
```

```python
roman_to_integer("XXXXX")
```

Output:
```
50
```

Create a function `integer_to_roman` that takes in an integer `n` then outputs the corresponding roman numeral string of the integer.

<!-- grade_id: cell-d9823aad8de8f225, points: 0, grade: false, solution: true -->
```python
def integer_to_roman(n):
    ### YOUR CODE HERE
    raise NotImplementedError
```

```python
integer_to_roman(40)
```

Output:
```
'XL'
```

<!-- grade_id: cell-0a42193d6c1b2b68, points: 20, grade: true, solution: false -->
```python
assert True
```