# Files, Modules, and Exceptions

References:

[1] B Downey, A. (2012). *Think Python: How to Think Like a Computer Scientist-2e.*

[2] Gries, P., Campbell, J., & Montojo, J. (2017). *Practical programming: an introduction to computer science using Python 3.6.* Pragmatic Bookshelf.

[3] Matthes, E. (2023). *Python crash course: A hands-on, project-based introduction to programming.*

[4] Downey, Allen, et al. "How to think like a computer scientist: learning with python 3." (2016). https://openbookproject.net/thinkcs/python/english3e

[4] `io` documentation - https://docs.python.org/3/library/io.html

[5] Errors and Exceptions - https://docs.python.org/3/tutorial/errors.html

Now that we have the foundation of the basic skills that we need to write programs, the next step is to make our programs more relevant and useable by being able to work with files - both reading and writing, and working with objects that can be read in a later operation (e.g. modules and pickle files). This allows our programs to be persistent - they keep at least some of their data in a permanent storage allowing us to use their results in a later time.

## 1 Files

Python has built-in file object that allow you to read and write text files. We can initialize such objects using the built-in `open` function.

In order, to open a file, you need to tell where that program is. By default, Python assumes that the file you want to read is in the same directory as the program that is doing the reading. As such, if our file is in another folder, we need to refer to it via absolute pathing or relative pathing.

```python
file = open('data/file_example.txt', 'r')
contents = file.read()
file.close()
print(contents)
```

Output:
```
First line of text
Second line of text
Third line of text
```

The second parameter in the `open` function specifies the mode in which you are opening the file. It can be:

| Mode | Value |
| ---  | ----- |
| Read | `r` | 
| Writing | `w` | 
| Exclusive Creation | `x` |
| Appending | `a` |
| Updating | `+` |
| Bytes | `b` |

We will demonstrate the use-case of each of the following mode in this notebook. But first, we discuss an important pattern when working with files, that is, using a context manager `with` statement.

### Using a `with` statement

A common programming pattern is to get access to a resource, do something with that resource, and then tidy up and release the resource. In our previous example, we read the file contents, then we tidied up by closing the file. However, if in between the reading and closing process, an error occurs, the `file.close()` will not be executed and the associated resources are never released.

Through the use of the `with` statement, we perform the `file.close()` process regardless of whether an error occurs. For this reason, the `with` statement is frequently used for file access.

```python
with open('data/file_example.txt', 'r') as file:
    contents = file.read()
print(contents)
```

Output:
```
First line of text
Second line of text
Third line of text
```

Under the hood, when we use the `with` statement, the objects's `__enter__` method is called throuh which the result of the call is assigned to the variable (in this case as `file`). After the block has been executed, the `__exit__` method of the object is called. In this case, the `file` object closes the file.

### Different modes when working with a file

When opening a file for writing or appending, the file will be created if it doesn't exists.

```python
with open('data/new_example.txt', 'w') as file:
    file.write('This is a new file.\n')

with open('data/new_example.txt', 'r') as file:
    contents = file.read()
print(contents)
```

Output:
```
This is a new file.
```

```python
with open('data/new_example.txt', 'a') as file:
    file.write('This is a new file.\n')

with open('data/new_example.txt', 'r') as file:
    contents = file.read()
print(contents)
```

Output:
```
This is a new file.
This is a new file.
```

`FileExistsError` will be raised if it already exists when opened for creating.

```python
# with open('data/new_example.txt', 'x') as file:
#     file.write('I want a new file.\n')
```

Opening a file for creating implies writing, thus, `x` behaves in a similar way as `w`.

```python
with open('data/exclusive_example.txt', 'x') as file:
    file.write('I want a new file.\n')
```

Output:
```
Cell In[6], line 1
----> 1 with open('data/exclusive_example.txt', 'x') as file:
      2     file.write('I want a new file.\n')

File ~/Library/Python/3.9/lib/python/site-packages/IPython/core/interactiveshell.py:310, in _modified_open(file, *args, **kwargs)
    303 if file in {0, 1, 2}:
    304     raise ValueError(
    305         f"IPython won't let you open fd={file} by default "
    306         "as it is likely to crash IPython. If you know what you are doing, "
    307         "you can use builtins' open."
    308     )
--> 310 return io_open(file, *args, **kwargs)

FileExistsError: [Errno 17] File exists: 'data/exclusive_example.txt'
```

```python
with open('data/exclusive_example.txt', 'r') as file:
    contents = file.read()
print(contents)
```

Output:
```
I want a new file.
```

We can add a `+` to allow for simultaneous reading and writing.

```python
with open('data/edited_example.txt', 'x') as file:
    file.write('This is a new file.\n')
```

Output:
```
Cell In[8], line 1
----> 1 with open('data/edited_example.txt', 'x') as file:
      2     file.write('This is a new file.\n')

File ~/Library/Python/3.9/lib/python/site-packages/IPython/core/interactiveshell.py:310, in _modified_open(file, *args, **kwargs)
    303 if file in {0, 1, 2}:
    304     raise ValueError(
    305         f"IPython won't let you open fd={file} by default "
    306         "as it is likely to crash IPython. If you know what you are doing, "
    307         "you can use builtins' open."
    308     )
--> 310 return io_open(file, *args, **kwargs)

FileExistsError: [Errno 17] File exists: 'data/edited_example.txt'
```

```python
with open('data/edited_example.txt', 'r+') as file:
    contents = file.read()
    file.write('I have read this file.\n')
print(contents)
```

Output:
```
I can read and edit this file.
I can read and edit this file.
```

```python
with open('data/edited_example.txt', 'r') as file:
    contents = file.read()
print(contents)
```

Output:
```
I can read and edit this file.
I can read and edit this file.
I have read this file.
```

```python
with open('data/edited_example.txt', 'w+') as file:
    contents = file.read()
    file.write('I can read and edit this file.\n')
print(contents)
```

```python
with open('data/edited_example.txt', 'r') as file:
    contents = file.read()
print(contents)
```

Output:
```
I can read and edit this file.
```

```python
with open('data/edited_example.txt', 'a+') as file:
    contents = file.read()
    file.write('I can read and edit this file.\n')
print(contents)
```

```python
with open('data/edited_example.txt', 'r') as file:
    contents = file.read()
print(contents)
```

Output:
```
I can read and edit this file.
I can read and edit this file.
```

```python
# !rm data/new_example.txt data/exclusive_example.txt data/edited_example.txt
```

### Using the `os` module

If we want to interface directly with the operating system and perform operations such as listing the directory contents, creating folders, and other os related operations. We can use the `os` module to do them in a portable manner.

```python
import os
```

For example, if we want to change the current directory, we can use the `os.chdir()` function.

```python
os.chdir('data')
```

```python
with open('file_example.txt', 'r') as file:
    contents = file.read()
print(contents)
```

Output:
```
First line of text
Second line of text
Third line of text
```

We can also list the directory using the `listdir` function. This creates a list containing the files in the specified directory.

```python
os.listdir('.')
```

Output:
```
['exclusive_example.txt',
 'new_example.txt',
 'sona.txt',
 'hopedale.txt',
 'file_example.txt',
 'edited_example.txt']
```

If we want to create folders, we can use the `mkdir` function.

```python
os.mkdir('some_folder')
```

```python
os.path.exists('some_folder')
```

Output:
```
True
```

To delete folders we can use `rmdir`

```python
os.rmdir('some_folder')
```

```python
os.path.exists('some_folder')
```

Output:
```
False
```

An alternative way of listing directories is using `glob`, which allows you to specify wildcards to get specific paths in a specified directory.

```python
from glob import glob
```

```python
glob('*')
```

Output:
```
['exclusive_example.txt',
 'new_example.txt',
 'sona.txt',
 'hopedale.txt',
 'file_example.txt',
 'edited_example.txt']
```

```python
glob('./*.txt')
```

Output:
```
['./exclusive_example.txt',
 './new_example.txt',
 './sona.txt',
 './hopedale.txt',
 './file_example.txt',
 './edited_example.txt']
```

```python
glob('*', include_hidden=True)
```

Output:
```
TypeError                                 Traceback (most recent call last)
Cell In[27], line 1
----> 1 glob('*', include_hidden=True)

TypeError: glob() got an unexpected keyword argument 'include_hidden'
```

### Techniques in reading files

When reading files, note that we always start at the current file cursor. This allows us to combine techniques that we want to use when reading files.

#### Reading the whole content or a specific number of characters.

When called with no arguments, `read` everything from the current file cursor all the way to the end of the file and moves the file cursor to the end of the file. With an argument, it reads that many characters and moves the file cursor after the characters that were just read.

```python
with open('file_example.txt', 'r') as file:
    first_ten_chars = file.read(10)
    the_rest = file.read()

print("The first 10 characters: ", first_ten_chars)
print("The rest of the file: ", the_rest)
```

Output:
```
The first 10 characters:  First line
The rest of the file:   of text
Second line of text
Third line of text
```

#### Making a `list` of lines from a file

If you want to et a Python list of strings containing the individual lines from a file, the function `readlines` work like `read`, but it splits up the lines into a list of strings.

```python
with open('file_example.txt', 'r') as file:
    lines = file.readlines()

print(lines)
```

Output:
```
['First line of text\n', 'Second line of text\n', 'Third line of text\n']
```

Notice that the line separators are not deleted upon the lien splitting. As such, we need to strip the string first of the leading spaces.

```python
with open('file_example.txt', 'r') as file:
    lines = file.readlines()

clean_lines = []
for line in lines:
    clean_lines.append(line.strip())
print(clean_lines)
```

Output:
```
['First line of text', 'Second line of text', 'Third line of text']
```

This can be rewritten using list comprehension.

```python
with open('file_example.txt', 'r') as file:
    lines = file.readlines()

clean_lines = [line.strip() for line in lines]
print(clean_lines)
```

Output:
```
['First line of text', 'Second line of text', 'Third line of text']
```

#### Reading line by line

When you want to do the same thing to every line from the file cursor to the end of a file, you can iterate the file using a `for` loop. On each iteration, the file cursor is moved to the beginning of the next line.

```python
with open('file_example.txt', 'r') as file:
    for line in file:
        print(line)
```

Output:
```
First line of text

Second line of text

Third line of text
```

```python
with open('file_example.txt', 'r') as file:
    lines = [line.strip() for line in file]

print(lines)
```

Output:
```
['First line of text', 'Second line of text', 'Third line of text']
```

#### Using different methods for different contexts

Sometimes we want to treat lines differently depending on context. For example, if the files that you're working on has a header section followed by a series of records, you can use `readline()` to read one line at a time.

In this example, we'll be using the number of colored fox fur pelts produced in Hopedale, Labrador in the years 1834-1842.

```python
with open('hopedale.txt', 'r') as file:
    # Skip the description line
    data = file.readlines()
data
```

Output:
```
['Coloured fox fur production, HOPEDALE, Labrador,, 1834-1925\n',
 '#Source: C. Elton (1942) "Voles, Mice and Lemmings", Oxford Univ. Press\n',
 '#Table 17, p.265--266\n',
 '      22   \n',
 '      29   \n',
 '       2   \n',
 '      16   \n',
 '      12   \n',
 '      35   \n',
 '       8   \n',
 '      83   \n',
 '     166   \n',
 '      62   \n',
 '      77   \n',
 '      42   \n',
 '      54   \n',
 '      56   \n',
 '       5   \n',
 '       9   \n',
 '      78   \n',
 '     151   \n',
 '      27   \n',
 '      14   \n',
 '      71   \n',
 '     261   \n',
 '      22   \n',
 '       3   \n',
 '      16   \n',
 '      70   \n',
 '       7   \n',
 '       2   \n',
 '      26   \n',
 '      27   \n',
 '      35   \n',
 '       9   \n',
 '      11   \n',
 '      73   \n',
 '      39   \n',
 '      22   \n',
 '      41   \n',
 '      31   \n',
 '      60   \n',
 '       7   \n',
 '      28   \n',
 '      43   \n',
 '      17   \n',
 '      26   \n',
 '      35   \n',
 '      60   \n',

```

We can use `readline()` to skip the header and the different comment lines using a `while` loop. Then later on counting all the number of fox pelts in the duration of the year.

```python
with open('hopedale.txt', 'r') as file:
    # Skip the description line
    file.readline()

    # Keep reading and skipping comment lines
    data = file.readline().strip()
    while data.startswith('#'):
        data = file.readline().strip()

    # This is the first piece of data
    total_pelts = int(data)

    # Read the rest of the data
    for data in file:
        total_pelts += int(data)
```

Instead of aggregating the data, we might just want to get all the values per year of the data and possibly store it in a list or store it in a text file for later use.

```python
with open('hopedale.txt', 'r') as file:
    # Skip the description line
    file.readline()

    # Keep reading and skipping comment lines
    data = file.readline().strip()
    while data.startswith('#'):
        data = file.readline().strip()

    # This is the first piece of data
    pelts_data = [int(data)]

    # Read the rest of the data
    for data in file:
        pelts_data += [int(data)]
```

```python
pelts_data
```

Output:
```
[22,
 29,
 2,
 16,
 12,
 35,
 8,
 83,
 166,
 62,
 77,
 42,
 54,
 56,
 5,
 9,
 78,
 151,
 27,
 14,
 71,
 261,
 22,
 3,
 16,
 70,
 7,
 2,
 26,
 27,
 35,
 9,
 11,
 73,
 39,
 22,
 41,
 31,
 60,
 7,
 28,
 43,
 17,
 26,
 35,
 60,
 75,
 18,
 58,
 38,
 28,
 26,
 29,
 43,
 16,
 40,
 25,
 28,
 12,
 22,
 13,
 52,
 90,
 103,
 15,
 26,
 29,
 102,
 31,
 9,
 33,
 90,
 95,
 17,
 35,
 59,
 115,
 12,
 26,
 38,
 202,
 25,
 83,
 92,
 58,
 0,
 19,
 63,
 108,
 25,
 38,
 231]
```

### Writing to a File

One of the simplest ways to save data is to write it to a file. We have already seen some ways to do this. The very basic of which is writing to an empty file.

```python
filename = 'programming.txt'

with open(filename, 'w') as file:
    file.write('I love programming.')
```

If we want to write multiple lines gievn a list, we can use the `writelines` method.

```python
lines_to_write = [
    'I love programming.',
    'I love games.'
]

with open(filename, 'w') as file:
    file.writelines(lines_to_write)
```

```python
with open(filename, 'r') as file:
    contents = file.read()
print(contents)
```

Output:
```
I love programming.I love games.
```

Notice that the `writelines` function does not automatically append a new line separator on each line. This has to be done manually. 

```python
lines_to_write = [line + "\n" for line in lines_to_write]
```

```python
with open(filename, 'w') as file:
    file.writelines(lines_to_write)
```

```python
with open(filename, 'r') as file:
    contents = file.read()
print(contents)
```

Output:
```
I love programming.
I love games.
```

### Persisting Python Objects

We now have learned one way to store data, that is thru text files. However, sometimes what we want is to directly store a data object such as a dictionary, function, or even machine learning models for future use then recreate them when we want to. This process is called *serialization* - which means converting a data object into a format that allows storage or transmission and recreation through a reverse process called *deserialization*.

The `pickle` module implements Python's procols for serializing and de-serializing a Python object. "Pickling" is process of serialization while "unpickling" is the deserialization.

Most of the Python objects we know can be pickled. (see [*What can be pickled and unpickled?*](https://docs.python.org/3/library/pickle.html#what-can-be-pickled-and-unpickled) for more info).

These includes:

- constants - `None`, `True`, `False`, `Ellipsis`, and `NotImplemented`;
- integers, floating point numbers, complex numbers;
- strings, bytes, bytearrays
- tuples, lists, sets, and dictionaries;
- functions (except for `lambda` functions)
- classes accessible from the top level of a module

To serialize a python object we use the `pickle.dump()` function on an opened file to be written as bytes. For example, if we want to serialize the list `pelts_data` from awhile ago and store it to a pickle file `data/pelts_data.pickle`, we do:

```python
pelts_data
```

Output:
```
[22,
 29,
 2,
 16,
 12,
 35,
 8,
 83,
 166,
 62,
 77,
 42,
 54,
 56,
 5,
 9,
 78,
 151,
 27,
 14,
 71,
 261,
 22,
 3,
 16,
 70,
 7,
 2,
 26,
 27,
 35,
 9,
 11,
 73,
 39,
 22,
 41,
 31,
 60,
 7,
 28,
 43,
 17,
 26,
 35,
 60,
 75,
 18,
 58,
 38,
 28,
 26,
 29,
 43,
 16,
 40,
 25,
 28,
 12,
 22,
 13,
 52,
 90,
 103,
 15,
 26,
 29,
 102,
 31,
 9,
 33,
 90,
 95,
 17,
 35,
 59,
 115,
 12,
 26,
 38,
 202,
 25,
 83,
 92,
 58,
 0,
 19,
 63,
 108,
 25,
 38,
 231]
```

```python
pelts_data[:10]
```

Output:
```
[22, 29, 2, 16, 12, 35, 8, 83, 166, 62]
```

```python
import pickle
```

```python
with open('pelts_data.pickle', 'wb') as file:
    pickle.dump(pelts_data, file)
```

```python
!ls
```

Output:
```
edited_example.txt    hopedale.txt          programming.txt
exclusive_example.txt new_example.txt       sona.txt
file_example.txt      pelts_data.pickle
```

To deserialize a pickled object, we use the `pickle.load()` function, then assign the value to a variable.

```python
with open('pelts_data.pickle', 'rb') as file:
    unpickled_pelts = pickle.load(file)
```

```python
unpickled_pelts == pelts_data
```

Output:
```
True
```

```python
unpickled_pelts[:10]
```

Output:
```
[22, 29, 2, 16, 12, 35, 8, 83, 166, 62]
```

## 2 Modules

Another type of file that we might want to create to make our programs persistent are *modules*. Modules are essentially `.py` files containing variables, function definitions, and statements which you can import into another Python program.

Writing modules not only help in making your program persistent but it allows for reusability and better maintainability as version control platforms such as `git` are much more suited to these type of files.

### Using modules

Up until now, we've already been using modules written by other people. For example, to gain access to the variables and functions of the `math` module, we've been importing it and then using the functions available to it.

```python
import math
```

You can use the `help` function to have an idea of the different variables, functions, and other classes which may be defined in that module.

```python
help(math)
```

For example, the math module contains the variable `pi` which is the number `pi` expressed up to the $10^{-15}$ decimal place.

```python
math.pi
```

We also gain access to functions such as the `sqrt` function which takes in the square root of a number.

```python
math.sqrt(9)
```

Sometimes, we want to import only what we need to use in our programs. Python let's you speify this by writing our `import` statement as:

```python
from math import sqrt, pi

radius = sqrt(9)
area = pi * radius**2
```

```python
area
```

### Defining your own modules

To save a code for later use, you can put it in a file with a `.py` extension. For example, the `example.py` contains functions that we have defined in previous session which is creating a celsius to fahrenheit temperature and vice versa converter.

```python
import example
```

Module files are essentially namespaces - a place where you are able to define names that will be visiable to the rest of your program.

```python
example.convert_to_fahrenheit(37)
```

```python
example.convert_to_celsius(98.6)
```

Notice that here, we did not specify any filepath to the module. This is because, Python by default has specific directories in which it searches for module. In particular, the `sys.path` variable contains the list of directories that the interpreter will search in for modules.

If a module is inside a directory, you need that directory to be included in the `sys.path` variable for you to be able to import it.

*Note: if you need to store several modules with similar functionalities in a directory, consider creating a package instead. This is out of the scope of this class but here's the reference for it https://docs.python.org/3/reference/import.html#regular-packages*

```python
import my_module
```

```python
os.getcwd()
```

```python
import sys
sys.path.append('../modules')
```

```python
import my_module
```

```python
my_module.my_function()
```

### Restoring or updating a module

Notice that if you change some function or variable definitions from a module, you cannot directly call them immediately.

```python
example.print_weather()
```

To reload or update a module, you can use the `importlib` library's function `reload` to use an updated version of your module.

```python
from importlib import reload
```

```python
reload(example)
```

```python
example.print_weather()
```

### What happens during an import

Notice that some other actions happen when we do import a module. When we imported the `example` module, there is an output being printed in addition to its namespace being available to us. This is because, any `.py` file is actually a program in itself and can be executed independently.

```python
os.chdir('..')
```

```python
!python example.py
```

But sometimes, we want to write code that should be only run when the module is directly executed not when the module is being imported. Python defines a special string variable `__name__` in every module to help us distinguish the two cases.

When a module is being imported, it's `__name__` is assigned as its module name. Meanwhile, if the module is being executed directly, the value assigned to `__name__` is `__main__`.

Notice the difference between

```python
reload(example);
```

and

```python
!python example.py
```

## 3 Exceptions

Whenever an error occurs during runtime, Python creates an *exception* object. If you write code that *handles* the exception, the program will continue running. If you don't handle the exception, the program will halt and show a *tracekback* which includes a report of the exception that was raised.

For example, dividing by zero creates an exception.

```python
10 * (1/0)
```

So does accessing an undefined variable.

```python
4 + spam*3
```

Or doing operations that are not supported by Python.

```python
'2' + 2
```

### Using try-except blocks

Sometimes we want to execute an operation that might cause an exception, but we don't want the program to stop. For example, if we're reading multiple files and want to process each and every file but unsure if each of the filename in our list actually exist. We can use a `try`-`except` block.

```python
try: 
    x = int(input("Please input an integer"))
except ValueError:
    print("Oops! That was not a valid integer")
```

Output:
```
Please input an integer 7.0
Oops! That was not a valid integer
```

```python
x = int(input("Please input an integer"))

```

Output:
```
Please input an integer 7.0
ValueError                                Traceback (most recent call last)
Cell In[61], line 1
----> 1 x = int(input("Please input an integer"))

ValueError: invalid literal for int() with base 10: '7.0'
```

```python
filepaths = ['data/file_example.txt', 'data/hopedale.txt',
             'data/non_existent_file.txt']
```

```python
for filepath in filepaths:
    try:
        with open(filepath, 'r') as file:
            contents = file.read()
        print(f"Successful reading of {filepath}.")
    except FileNotFoundError:
        print(f"File {filepath} does not exist.")
```

### Raising our own exceptions

If we want to deliberately raise exceptions because we deem such cases as an error for our program. We can `raise` an Exception. For a list of all possible exception and their corresponding hierarchy, please look at the documentation of it [here](https://docs.python.org/3/library/exceptions.html#exception-hierarchy).

```python
raise ValueError
```

```python
raise ValueError("Invalid value for the variable")
```

For example, if we're creating a function that checks the age of a person, we don't expect it to have a negative value. As such, we may raise an exception for it.

```python
def get_age():
    age = int(input("Please enter your age: "))
    if age < 0:
        # Create an instance of exception
        age_error = ValueError(f"{age} is not a valid age.")
        raise age_error
    return age
```

```python
get_age()
```

### Multiple except clause

A `try` statement may have more than one except clause to specify the different handlers for different exceptions.

```python
try:
    raise Exception
except ValueError:
    print("There has been a value error")
except ZeroDivisionError:
    print("There has been a zero divison error")
```

### Accessing the argument of the exception

The *except* clause may specify a variable after the exception name. This variable is bound to the exception instance which typically has an `args` attribute that stores the arguments. You can access this arguments and use their value in your code for handling the exception.

```python
try:
    raise Exception("There is an error", "It is an exception")
except Exception as err:
    print(err)
    print(type(err))
    print(err.args)
    x, y = err.args
    print(x)
    print(y)
```

### The `else` block

We can also add an `else` block after the `try-except` statement. In this case, when present, it must follow all the `except` clauses. The `else` block works like this: Python attempts to run the code in the `try` block. If the `try` block was successful, the code goes in the `else` block. The `except` block tells Python what to do in case a certain exception arises when it tries to run the code in the `try` block.

For example we have the function `number_of_lines()` which takes in the a `filepath` then returns the number of lines in that file if successfully read. If not, it will print the message `"Cannot open the file [filepath]."`.

```python
def number_of_lines(filepath):
    try:
        file = open(filepath, 'r')
    except OSError:
        print(f"Cannot open the file {filepath}.")
    else:
        return len(file.readlines())
```

```python
number_of_lines('data/file_example.txt')
```

### Defining clean-up actions

If we want to perform actions that will be excecuted regardless as to whether an error has been raised or not. We can use the `finally` clause of the `try` statement.

When the `finally` clause is present, the finally clause will execute as the last task before the `try` statement completes.

The `try` statement has another optional clause which is intended to define clean-up actions that will be executed under all circumstances. 

```python
try:
    raise Exception
except ValueError:
    print("There has been a value error")
except ZeroDivisionError:
    print("There has been a zero divison error")
finally:
    print("Goodbye, World!")
```

Note the order in which the finally statement is evaluated - if a `try` statement reaches  a `break`, `continue`, or `return` statement, the `finally` clause will execute just prior to the break, continue, or return statement.

```python
def bool_return():
    try:
        return True
    finally:
        return False
```

```python
bool_return()
```

## 4 Hands-on Exercises

### Number Pairs

Create a function named `create_number_pairs()` with a required parameters, `n` and `out_filepath` that generates a pair of `n` numbers between `0` and `100` then writes the pairs into a text file at filepath `out_filepath` with each number separated by a comma `,` then each pair is written per line of the text file.

<!-- grade_id: cell-f76f57d7956b69e4, points: 0, grade: false, solution: true -->
```python
from random import random


def create_number_pairs(n, out_filepath):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Create a function named `create_sum_number_pairs()` which has two required parameters `in_filepath` and `out_filepath`, and one optional parameter `header` which defaults to `False`. The function will read the pair of numbers present in `in_filepath` then writes each line on `out_filepath` corresponding to the first number, second number, and their sum separated each separated by `,`. When the optional parameter `header` is set to true, it writes the header `x,y,x+y` on top of the output file.

<!-- grade_id: cell-bcf13806da2fcd77, points: 0, grade: false, solution: true -->
```python
def create_sum_number_pairs(in_filepath, out_filepath, header=False):
    # YOUR CODE HERE
    raise NotImplementedError()
```

### Word Tools

Create a module neamed `wordtools.py` that contains the following function:

1. `read_file(filepath, header=False)` which reads the text file in the `filepath` then returns a string containing all of its contents. If the `header` parameter is set to `True`, it will skip the first lines that start with `#`.
2. `clean_text(s, text, word=False)` which takes in a string of characters `s` then removes each of the character in `s` on the input `text`. The cleaned texted is return as the value of this function. If the `word` parameter is set to true, it instead removes all the instances of the string `s` in the `text`.
3. `remove_period(text)` which takes in a string `text` then removes the `.` character if it occurs at the end of the string, it returns the original text if `.` does not occur at the end of the string. The cleaned text is returned as the value of this function.

### SONA: Most Spoken Words

You are given a text file (found at `data/sona.txt`) containing the transcript of the state of the nation address of the president last July 24, 2023. In this exercise, we will read this text file then get the most spoken words based on this transcript.

Create a function named `sona_word_counts(filepath)` which takes in the `filepath` of a sona transcript text then returns a dict-like object that shows the word frequencies occuring on that text. The keys of the returned dict-like object are the words, then the value is the frequency count of that word in the text.

Before obtaining the word counts, perform a text cleaning by doing these operations:

1. Skip the headers which are starting lines that start with `#`
2. Remove special characters such as `,;'!-`
3. Remove occurrences of the `[applause]`
4. Remove `.` only if they are used as sentence enders.
5. Case normalize each word (take their lowercase).

Use the functions that you defined in `wordtools.py` to help you in this exercise.

<!-- grade_id: cell-86fa425f3ce34a1f, points: 0, grade: false, solution: true -->
```python
from collections import Counter
import wordtools


def sona_word_counts(filepath):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
wc = sona_word_counts('data/sona.txt')
```

What is the top 20 common words in the speech of the president last SONA?

<!-- grade_id: cell-d521adcf2166756e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Positive Vibes Only

Write a function named `read_pos_int` which uses the `input` dialog to prompt the user for a positive integer and then checks the input to confirm that it meets the requirements. Raise a `ValueError` if the input is a negative `int` with the message `"Negative int not allowed!"`. For all other cases, handle the error. Specifically, if the user does not input anything, print a message `"Input something."`. If the input cannot be converted to `int`, handle the error and print the message `"Input cannot be interpreted as an integer."`.

<!-- grade_id: cell-ab2edae3ed39c102, points: 0, grade: false, solution: true -->
```python
def read_pos_int():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d2170458418f873a, points: 20, grade: true, solution: false -->
```python
assert True
```