# Object Oriented Programming

References:

[1] B Downey, A. (2012). *Think Python: How to Think Like a Computer Scientist-2e.*

[2] Gries, P., Campbell, J., & Montojo, J. (2017). *Practical programming: an introduction to computer science using Python 3.6.* Pragmatic Bookshelf.

[3] Matthes, E. (2023). *Python crash course: A hands-on, project-based introduction to programming.*

[4] Downey, Allen, et al. "How to think like a computer scientist: learning with python 3." (2016). https://openbookproject.net/thinkcs/python/english3e

Python is an object-oriented programming language, which means that it provides features that support **object-oriented proramming**.

Up to now, most of the programs we have been writing use a procedural programming paradigm, in which, the focus is to write functions or procedures which operate on data. In OOP, the focus is on the creation of **objects** which contain both data and functionality together.

Usually each object is defined according to some object or concept in the real world, and the functions that operate on that object correspond to the ways real-world objects interact.

## 1 Classes and Objects

### Programmer-defined types

```python
a = [4, 5, 1, 2, 3, 0, 5]
```

```python
print(a)
```

```python
a.sort()
print(a)
```

A programmer-defined type is also called a **class**. A class definition looks like this.

```python
class Point:
    """Point class represents and manipulates x, y coordinates"""
    def __init__(self, x, y):
        self.x = x
        self.y = y
```

Defining a class named `Point` creates a **class object**.

```python
Point
```

The class onbject is like a factory for creating objects. To create a Point, you call `Point` as if it were a function.

```python
blank = Point()
blank
```

```python
print(blank)
```

```python
blank.x
```

Creating a new object is called **instantiation**, and the object is an **instance** of the class.

### Attributes

You can assign values to an instance using dot notation.

```python
blank.x = 3.0
blank.y = 4.0
```

```python
blank.x
```

```python
blank.y
```

You can pass an instance as an argument in a function

```python
def print_point(p):
    print(f'({p.x}, {p.y})')
```

```python
print_point(blank)
```

Every class should have a method with the special name `__init__`. This initializer method is automatically called whenever a new instance of the class is created. It gives the programmer the opportunity to setup the attributes required within the new instance by giving them their initial state/values.

```python
class Point:
    """Point class represents and manipulates x, y coordinates"""

    def __init__(self, x=0, y=0):
        """Initialize a new point at x,y"""
        self.x = x
        self.y = y
        self.p = Momentum()
```

```python
class Momentum:
    """Point class represents and manipulates x, y coordinates"""

    def __init__(self, x=0, y=0):
        """Initialize a new point at x,y"""
        self.px = x/2
        self.py = y/2
```

```python
test = Point()
print(test.x, test.p.px)

test.x = 5
test.y = 5
print(test.x, test.p.px)
```

```python
test.p.px = 5
```

```python
test.p.px
```

### Adding other methods

```python
class Point:
    """Point class represents and manipulates x, y coordinates"""

    def __init__(self, x=0, y=0):
        """Initialize a new point at x,y"""
        self.x = x
        self.y = y

    def distance_from_origin(self):
        """Compute the distance of point to origin"""
        return ((self.x ** 2) + (self.y ** 2))**0.5
```

```python
p = Point(3, 4)
q = Point(5, 12)
```

```python
p.x, p.y
```

```python
q.x, q.y
```

```python
p.distance_from_origin()
```

```python
q.distance_from_origin()
```

Notice that we usually have to print a specific objec and expect it to return something meaningful. However, with our current definition, this is still not possible.

```python
str(p)
```

```python
print(p)
```

We can define a method `__str__` to introduce our intended behavior.

```python
class Point:
    """Point class represents and manipulates x, y coordinates"""

    def __init__(self, x=0, y=0):
        """Initialize a new point at x,y"""
        self.x = x
        self.y = y

    def __str__(self):
        return f"(x is {self.x}, y is {self.y})"

    def distance_from_origin(self):
        """Compute the distance of point to origin"""
        return ((self.x ** 2) + (self.y ** 2))**0.5
```

```python
p = Point(3, 4)
```

```python
print(p)
```

```python
str(p)
```

### Instances as return values

Functions and methods can return instances. For example, given two `Point` objects, find their midpoint.

```python
def midpoint(p1, p2):
    """Return the midpoint of points p1 and p2"""
    mx = (p1.x + p2.x) / 2
    my = (p1.y + p2.y) / 2

    return Point(mx, my)
```

```python
p = Point(3, 4)
q = Point(5, 12)
r = midpoint(p, q)
```

```python
r
```

```python
print(r)
```

We can do this as a method instead. For example, we wish to find the midpoint halfway between it and some other target point.

```python
class Point:
    """Point class represents and manipulates x, y coordinates"""

    def __init__(self, x=0, y=0):
        """Initialize a new point at x,y"""
        self.x = x
        self.y = y

    def __str__(self):
        return f"({self.x}, {self.y})"

    def halfway(self, target):
        """Return the halway point between self and the target"""
        mx = (self.x + target.x) / 2
        my = (self.y + target.y) / 2

        return Point(mx, my)

    def distance_from_origin(self):
        """Compute the distance of point to origin"""
        return ((self.x ** 2) + (self.y ** 2))**0.5
```

```python
p = Point(3, 4)
q = Point(5, 12)
r = p.halfway(q)
r
```

```python
print(r)
```

### Objects are mutable

```python
print(p)
```

```python
p.x = 10
p.y = 10
```

```python
print(p)
```

```python
class Point:
    """Point class represents and manipulates x, y coordinates"""

    def __init__(self, x=0, y=0):
        """Initialize a new point at x,y"""
        self.x = x
        self.y = y

    def __str__(self):
        return f"({self.x}, {self.y})"

    def halfway(self, target):
        """Return the halway point between self and the target"""
        mx = (self.x + target.x) / 2
        my = (self.y + target.y) / 2

        return Point(mx, my)

    def distance_from_origin(self):
        """Compute the distance of point to origin"""
        return ((self.x ** 2) + (self.y ** 2))**0.5

    def set_position(self, x, y):
        """Set the position of the point to a new x and y"""
        self.x = x
        self.y = y

    def move(self, dx, dy):
        """Move the object by the deltas"""
        self.x += dx
        self.y += dy
```

```python
p = Point()
```

```python
print(p)
```

```python
p.set_position(3, 4)
```

```python
p.x = 100
```

```python
print(p)
```

```python
p.move(7, 6)
```

```python
print(p)
```

```python
import numpy as np
```

```python

```

## 2 Polymorphism

Some languages, including Python, make it possible to have different meanings for the same operator when applied to different types. For example, `+` in Python means quite different things for integers and for strings. This feature is called **operator overloading**.

In our `Point` object, we can define how `+` operates by including an `__add__` method.

```python
4+4
```

```python
"4"+"4"

```

```python
class Point:
    """Point class represents and manipulates x, y coordinates"""

    def __init__(self, x=0, y=0):
        """Initialize a new point at x,y"""
        self.x = x
        self.y = y

    def __str__(self):
        return f"({self.x}, {self.y})"

    def __add__(self, other):
        return Point(self.x + other.x, self.y + other.y)

    def halfway(self, target):
        """Return the halway point between self and the target"""
        mx = (self.x + target.x) / 2
        my = (self.y + target.y) / 2

        return Point(mx, my)

    def distance_from_origin(self):
        """Compute the distance of point to origin"""
        return ((self.x ** 2) + (self.y ** 2))**0.5

    def set_position(self, x, y):
        """Set the position of the point to a new x and y"""
        self.x = x
        self.y = y

    def move(self, dx, dy):
        """Move the object by the deltas"""
        self.x += dx
        self.y += dy
```

```python
p = Point(3, 4)
q = Point(5, 12)
r = p + q
```

```python
print(r)
```

We can do the same for multiplication.

```python
class Point:
    """Point class represents and manipulates x, y coordinates"""

    def __init__(self, x=0, y=0):
        """Initialize a new point at x,y"""
        self.x = x
        self.y = y

    def __str__(self):
        return f"({self.x}, {self.y})"

    def __add__(self, other):
        return Point(self.x + other.x, self.y + other.y)

    def __mul__(self, other):
        return self.x * other.x + self.y * other.y

    def __rmul__(self, other):
        return Point(other * self.x, other * self.y)

    def halfway(self, target):
        """Return the halway point between self and the target"""
        mx = (self.x + target.x) / 2
        my = (self.y + target.y) / 2

        return Point(mx, my)

    def distance_from_origin(self):
        """Compute the distance of point to origin"""
        return ((self.x ** 2) + (self.y ** 2))**0.5

    def set_position(self, x, y):
        """Set the position of the point to a new x and y"""
        self.x = x
        self.y = y

    def move(self, dx, dy):
        """Move the object by the deltas"""
        self.x += dx
        self.y += dy
```

```python
p = Point(3, 4)
q = Point(5, 12)
p * q
```

```python
r = 3 * p
print(r)
```

```python
p * 3
```

Most methods only work for a specific type, but there are certain operations that you want to apply to many types, such as arithmetic operations. If many types support the same set of operations, we can write functions that work on any of those types.

For example, consider `multadd` operation.

```python
def multadd(x, y, z):
    return x*y + z
```

This function will work for any values of x and y that can be multipleied and for any value of `z` that can be added to the product

```python
multadd(3, 2, 1)
```

```python
multadd(1, p, q)
```

```python
multadd(p, q, 2)
```

A function like this that can take arguments with different types is called **polymorphic**.

To determine whether a function can be applied to a new type, we apply Python's fundamental rule of polymorphism, called the **duck typeing rule**: *If all operations inside the function can be applied to the type, the function can be applied to the type.*

## 3 Inheritance

Giving one class the same methods as another is one way to make them polymorphic, but it suffers from repeated code if we always do that. If a programmer forgets just one line of code, the whole program can fail for reasons that will be difficult to track down.

A better approach is to use a feature called *inheritance*, which allows you to recycle code in yet another way.

Whenever you create a class using inheritance, you new class automatically inherits all of the attributs of the class *object*, much like a child inherits attributes from their parents.

For example, let's say we are managing people at a university. There are students and faculty. Both students and faculty have names, postal addresss, and email address; each student also has a student number, a list of course taken, and a list of courses he or she is currently taking. Each faculty member has a faculty number and a list of courses he or she is currently teaching.

We can define a `Faculty` and a `Student` class such that both of them are subclasses of a `Member` class which signifies that they are both members of a university.

```python
class Member:
    """A member of a university"""

    def __init__(self, name, address, email):
        self.name = name
        self.address = address
        self.email = email


class Faculty(Member):
    """A faculty member at a university"""

    def __init__(self, name, address, email, faculty_num):
        super().__init__(name, address, email)
        self.faculty_number = faculty_num
        self.courses_teaching = []


class Student(Member):
    """A student member at a university"""

    def __init__(self, name, address, email, student_num):
        super().__init__(name, address, email)
        self.student_number = student_num
        self.courses_taken = []
        self.courses_taking = []
```

```python
faculty_leo = Faculty('Leodegario Lorenzo II', 'Quezon City',
                      'llorenzo@aim.edu', 221039)
student_leo = Student('Leodegario Lorenzo II', 'Quezon City',
                      'llorenzo@aim.edu', 708472)
```

```python
faculty_leo.name, faculty_leo.address, faculty_leo.email
```

```python
student_leo.name, student_leo.address, student_leo.email
```

```python
faculty_leo.courses_teaching.extend(['PDS', 'ML2', 'Calculus'])
```

```python
faculty_leo.faculty_number, faculty_leo.courses_teaching
```

```python
student_leo.courses_taking.extend(['GC', 'SRP'])
```

```python
student_leo.student_number, student_leo.courses_taking
```

## 4 Hands-on Exercises

### Today is a Good Day to Dice

X Make a class `Dice` 
    
X    with one attribute called `sides`, 
        
X        which has a default value of `6`. 

X Write a method called `roll_die()` 
    
    that prints a random number between 1 and the number of side the die has.

<!-- grade_id: cell-6daba71c63f59744, points: 0, grade: false, solution: true -->
```python
import random

class Dice:
    def __init__(self, sides=6):
        self.sides = sides

    def roll_die(self):
        return random.randrange(1, self.sides+1, 1)
```

Make a 6-sided die 

    and roll it `10` times. 

Place the result in a list named `die_result`.

```python
n_roll = 100
sides = 20

# WORKING CODE
d = Dice(sides)
die_result = []
for i in range(n_roll):
    die_result.append(d.roll_die())

# TEST CODE
#print(die_result)
print('DEBUG - number of rolls:', len(die_result))
print('DEBUG - max roll value:', max(die_result))
print('DEBUG - min roll value:', min(die_result))
```

```python

```

<!-- grade_id: cell-aaa3a942919e9f8c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
pass
### END SOLUTION
```

Make a 10-sided die and roll it `10` times. Place the result in a list named `die_result`.

### Two Points Determine a Line

Consider this code:

```
>>> segment = LineSegment(Point(1, 1), Point(3, 2))
>>> segment.slope()
0.50
>>> segment.length()
2.23606797749979
```

Write two classes `Point` and `LineSegment` so that you can run the above code and get the same results.

1. Write a `Point` class with `__init__` method that takes two numbers as parameters.
2. Write a `LineSegment` class whose constructor takes two `Points` as parameters. The first `Point` should be the start of the segment.
3. Write a `slope` method in the class `LineSegment` that computes the slope of the segment.
4. Write a `length` method in `LineSegment` that computes the length of the segment.

<!-- grade_id: cell-24bd51b35048aad6, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
pass
### END SOLUTION
```

## Card and Deck of Cards

1. Create a `Card` class with the following attributes:

- `rank` (e.g., Ace, 2, 3, ..., King)
- `suit` (e.g., Hearts, Diamonds, Clubs, Spades)

2. Implement a `__str__` method in the Card class to display the card's rank and suit.

3. Create a `Deck` class with the following methods:

- `__init__`: Initializes an empty list to represent the deck.
- `populate_deck`: Populates the deck with 52 standard playing cards.
- `shuffle`: Shuffles the deck using the random.shuffle function.
- `draw_card`: Removes and returns the top card from the deck.
- Implement a `__str__` method in the Deck class to display the entire deck.

4. Create an instance of the `Deck` class, populate it, shuffle it, and then draw and display a few cards.

<!-- grade_id: cell-e480920840c9bf0b, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
pass
### END SOLUTION
```

<!-- grade_id: cell-5e2d61c76e1dd42c, points: 20, grade: true, solution: false -->
```python
assert True
```